# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Pre-Lab #3**<span style="color:red"> $\to$ (7 possible marks)</span> </span></div> 
# Electric Potential & Electric Fields

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);">Understand the term *field* as it is used in physics, and distinguish between scalar fields and vector fields.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Become familiar with different ways to represent fields in two-dimensional plots.</span></b>

The term field comes up a lot in physics: you've probably heard the terms *gravitational field*, *electric field*, *force field*, and *magnetic field*.  But what do we mean when we call something a field?  Mathematically speaking, a **field** is any quantity $\varphi(x, y, z, t)$ that has a value associated with all points $(x,y,z)$ in space and moments $t$ in time.  The type of value depends on the type of field: $\varphi$ might be **scalar**-valued or **vector**-valued.  

This definition might be unfamiliar at first, or it might seem simple but too abstract to be meaningful.  Through a series of examples and questions, this pre-lab will serve as a gentle introduction to scalar- and vector-valued fields and their representations.  We'll conclude by revisiting the formal definition.


***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/images/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Pre-Lab 1.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
```

## Scalar Fields

Suppose you were to walk over rough terrain in a perfectly straight line, measuring your altitude $h$ above sea level at regular intervals.  If you were to denote distance from the starting point by the variable $x$, you could define a function $h(x)$.  Plotting this function produces something that looks like a cross-section of the terrain:

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/y0_hvsx_fill.png" alt="image info" style="width: 45%; height: auto;" /><br>
<b>Fig. 1: A plot of the height above sea level as a function of $x$.</b></center>
</p>

But what if we wanted to plot the altitude over a *patch* of land?  We might define a second variable $y$, draw an imaginary "grid" over the patch of land, and measure the altitude at different gridpoints.  By doing so, we define a function $h(x,y)$.  The more closely spaced our measurements, the more detailed our map.  In Fig. 1 we've labeled the vertical axis of our graph as $h(x,0)$ in anticipation of this.

But how should we represent our function $h(x,y)$?  Well... we could construct a 3-dimensional physical model of the land, maybe.  But what if we wanted soemthing that we could print out on a map? Consider the topographic map shown in Fig. 2.

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/cp.png" alt="image info" style="width: 95%; height: auto;" /><br>
<b>Fig. 2: A topographic map of the rock quarry showing the elevation above sea level as a function of $(x, y)$ coordinates. The difference in height between adjacent contours is $2\rm\ m$, and every second contour line is labeled with its height above sea level in meters. </b></center>
</p>

This type of plot is likely already familiar to you.  It's constructed by drawing lines of constant altitude $h$ &ndash; these lines are called **contours**, and accordingly the plot in Fig. 2. is called a **contour plot**.  The plot above specifies the $x$ position with letters, which is something we wouldn't be inclined to do, but that's okay - if you follow with your finger, you can read from the map above that $h(T, 6) \approx 24 \rm\ m$ above sea level.

***
**<span style="color:blue">Question 1:</span>**  **<span style="color:red">(1 mark)</span>**

Consider the orange arrow in Fig. 2.  If you were to walk along that arrow, would you be

(a) Ascending <br>
(b) Descending <br>
(c) Neither

***
**<span style="color:blue">Answer 1:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a1 = ...
```

```python
grader.check("Q1")
```

***
**<span style="color:blue">Question 2:</span>**  **<span style="color:red">(1 mark)</span>**

Consider the blue arrow in Fig. 2.  If you were to walk along that arrow, would you be

(a) Ascending <br>
(b) Descending <br>
(c) Neither

***
**<span style="color:blue">Answer 2:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a2 = ...
```

```python
grader.check("Q2")
```

***
**<span style="color:blue">Question 3:</span>** **<span style="color:red">(1 mark)</span>** 

Consider the red arrow in Fig. 2.  If you were to walk along that arrow, would you be

(a) Ascending <br>
(b) Descending <br>
(c) Neither

***
**<span style="color:blue">Answer 3:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a3 = ...
```

```python
grader.check("Q3")
```

***
**<span style="color:blue">Question 4:</span>**  **<span style="color:red">(1 mark)</span>**

Along which arrow in Fig. 2 is the terrain the steepest?

(a) Orange arrow <br>
(b) Blue arrow <br>
(c) Red arrow


***
**<span style="color:blue">Answer 4:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4 = ...
```

```python
grader.check("Q4")
```

<br>

***


Alternatively, we could encode this data with a colour scale.  Figure 3 below encodes the same information as Figure 2 by representing low altitudes with the colour blue and high altitudes with the colour red.  This is sometimes referred to as a **false colour plot**.

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/topo_fill.png" alt="image info" style="width: 85%; height: auto;" /><br>
<b>Fig. 3: A false colour plot of the height above sea level as a function of $x$ and $y$.</b></center>
</p>

Representative contour lines are often superimposed on false colour plots as guides to the eye.  Other than the contours of constant height, Fig. 4 is identical to Fig. 5.

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/topoContour_fill.png" alt="image info" style="width: 85%; height: auto;" /><br>
<b>Fig. 4: A false colour plot of the height above sea level as a function of $x$ and $y$ with some representative contour lines included.</b></center>
</p>

Elevation isn't the only thing that can be represented by contour plots and false colour plots - have a look at the infrared image in Figure 5.  The colour scale represents temperature in degrees Fahrenheit ($45.8^\circ\rm F = 7.7^\circ\rm C$ and $78.2^\circ\rm F = 25.7^\circ\rm C$).     
<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/deer.jpg" alt="image info" style="width: 40%; height: auto;" /><br>
<b>Fig. 5: An infrared image of a deer in the forest.  Image obtained from <a href = "https://commons.wikimedia.org/wiki/File:A_deer_in_the_woods._It_was_pitch_black,_but_this_camera_does_not_use_ANY_visible_light._FLIR_infrared._%286968297826%29.jpg">Wikimedia Commons</a>.</b></center>
</p>

All of the plots that we've consider so far are used to represent **scalar fields**.  Recall that a **scalar** is a quantity that has a numerical value, but no associated direction.  Simply put - a scalar is just a number.  Height above sea level in a rock quarry and temperature in the forest are both examples of scalar fields.  

***
## Vector Fields

A **vector** is a quantity that has both *magnitude* and *direction*.  Vectors are typically represented by arrows whose lengths correspond to their magnitude.  A **vector field** is a quantity that has a vector value for *every* point in space and time.  Think about the air flowing around you right now, for instance.  At each point, the air has some speed and some direction.  

When representing vector fields pictorially we obviously can't draw *every* vector, as there are infinitely many of them.  Usually, we'll draw representative arrows throughout a region of space, as in Figure 6.  Fig. 6(a) represents a cross-section of the gravitational field around the Earth-moon system and (b) represents the velocity of the water at the surface of a stream.  
<br>
<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/vectorField.png" alt="image info" style="width: 70%; height: auto;" /><br>
<b>Fig. 6: Examples of vector fields. (a) the gravitational field in the space around the Earth and moon. (b) The surface water velocity in a river containing a large rock.  Images obtained from <a href = "https://math.libretexts.org/Bookshelves/Calculus/Book%3A_Calculus_%28OpenStax%29/16%3A_Vector_Calculus/16.1%3A_Vector_Fields">LibreTexts</a>.</b></center>
</p>

Notice how colour has been used in both plots to emphasize the magnitude of the vectors.  This due to the fact that representing a vector's magnitude by length alone can make a plot hard to read.  But the magnitude of a vector is just a number - that is, it's a scalar quantity!  In Fig. 6(b), a false colour plot represents the speed of the water at different points in space, with red representing high speeds and blue representing low speeds.  You might say that the velocity field is a vector field, but the "speed field" is a scalar field.

***
**<span style="color:blue">Question 5:</span>**  **<span style="color:red">(1 mark)</span>**

Which of the following is an example of a vector field?

(a) The temperature in your room <br>
(b) The flow of heat from a radiator in your room <br>
(c) Neither

***
**<span style="color:blue">Answer 5:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a5 = ...
```

```python
grader.check("Q5")
```

***
**<span style="color:blue">Question 6:</span>**  **<span style="color:red">(1 mark)</span>**

Which of the following is an example of a scalar field?

(a) The motion of ammonia ice through the atmosphere of Jupiter <br>
(b) The concentration of ammonia throughout the atmosphere of Jupiter <br>
(c) Neither


***
**<span style="color:blue">Answer 6:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a6 = ...
```

```python
grader.check("Q6")
```

<br>

***
## Calculating a Vector Field from a Scalar Field

These last few examples emphasize the idea that vector fields often have associated scalar fields, and vice versa.  In many cases it is possible to calculate a vector field from a scalar field via the **gradient** of a scalar field, which describes both the *direction* and *rate* of fastest increase of a field for all points in space.  Since it has a direction, the gradient of a scalar field $\varphi$ is clearly a vector quantity, and is denoted by $\nabla \varphi$.  The resulting field is called a **gradient field**. 

As an example, from the topographic map of Fig. 4, we can generate a vector field that characterizes the steepest ascent at each point:

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab3/images/topoVector_fill.png" alt="image info" style="width: 85%; height: auto;" /><br>
<b>Fig. 7: A false colour plot of the height above sea level as a function of $x$ and $y$ with some representative contour lines and the steepness vector field.</b></center>
</p>

We can do this by calculating the quantities $dh/dx$ and $dh/dy$ which tells us how quickly $h$ changes in the $x$ and $y$ directions, respectively.  If, at a particular point in space, $dh/dx$ is large and $dh/hy$ is small, then the hill is steep in the $x$ direction and shallow in the $y$-direction.  If we were to construct the vector
<br>

\begin{align}
\mathbf{s} \equiv \mathbf{\nabla} h = \frac{dh}{dx}\,\hat\imath + \frac{dh}{dy}\,\hat\jmath,
\end{align}
<br>
it would tell us the direction of steepest ascent at that point on the mountain.  Fortunately, we don't have to do this calculation manually.  We can use some built-in Python functions to do the calculations for us.  Notice that, in this case, $\mathbf{s}$ is always perpendicular to the contour lines of constant elevation.  

An important example of gradient fields in physics are **conservative** force fields.  These are the gradient fields of the potential energy $U$:
<br>

\begin{align}
\mathbf{F}(x,y,z) = -\mathbf{\nabla}U \equiv -\left(\frac{dU}{dx}\,\hat\imath + \frac{dU}{dy}\,\hat\jmath + \frac{dU}{dy}\,\hat{k}\right)
\end{align}


***
**<span style="color:blue">Question 7:</span>**  **<span style="color:red">(1 mark)</span>**

Notice the negative sign in the equation for $\mathbf{F}$. Considering our interpretation of the field $\mathbf{s} = \mathbf{\nabla} h$, which of the following statements is true?

(a) $\mathbf{F}$ points in the direction of fastest *increase* of the potential energy $U$ <br>
(b) $\mathbf{F}$ points in the direction of fastest *decrease* of the potential energy $U$ <br>
(c) Neither

***
**<span style="color:blue">Answer 7:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, or c) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a7 = ...
```

```python
grader.check("Q7")
```

***

Similarly, time-invariant electric fields $\mathbf{E}$ are negative gradient fields of the electric potential $V$. In this week's lab, you will be measuring a two-dimensional electric potential $V(x,y)$.  From the data you collect, you will find the electric field $\mathbf{E}(x,y)$ via the equation
<br>

\begin{align}
\mathbf{E}(x,y) = -\left( \frac{dV}{dx}\,\hat\imath + \frac{dV}{dy}\,\hat\jmath \right)
\end{align}

***

Recall our formal definition now: a **field** is *any* quantity $\varphi(x, y, z, t)$ that has a value associated with all points $(x,y,z)$ in space and moments $t$ in time.  If the field is time-invariant, or if we're just looking at a "snapshot" of the field, we might consider only its values in space and write $\varphi(x, y, z)$.  Moreover, if we're only considering values over a two-dimensional surface, we might describe our field by $\varphi(x, y)$.  

The abstraction of mathematics equips us with a concrete set of tools we can use for analysis and interpretation, regardless of what type of field we're dealing with.  Learning to work with these mathematical ideas takes some practice, but even if you find yourself lost in the abstraction you should never lose sight of the *physics* behind the mathematics.  Fields are a fundamentally physically motivated concept which find *many* uses in physics and beyond.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

**<span style="color:blue">Feedback Question:</span>**  

We welcome your feedback on the PHYS 121 pre-labs!  Please feel free to include any comments you have about this pre-lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs and pre-labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the pre-lab, ...

This question is optional and will have no impact on your pre-lab or lab grade.

***
**<span style="color:blue">Feedback Answer:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

***

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Lab Canvas gradebook.

Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/phys231/Python/images/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export()
```

***
Last update: January 26, 2023

Tests for Q1 (points: None)
```
>>> a1.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.


Tests for Q2 (points: None)
```
>>> a2.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.


Tests for Q3 (points: None)
```
>>> a3.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.


Tests for Q4 (points: None)
```
>>> a4.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.


Tests for Q5 (points: None)
```
>>> a5.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.


Tests for Q6 (points: None)
```
>>> a6.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.


Tests for Q7 (points: None)
```
>>> a7.lower() in ["a", "b", "c"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', or 'c'.
